---
# theme
# “bootstrap”, “cerulean”, “cosmo”, “darkly”, “flatly”, “journal”, “lumen”, “paper”, 
# “readable”, “sandstone”, “simplex”, “spacelab”, “united”, “yeti”
# highlight: `default`, `tango`, `pygments`, `kate`, `monochrome`, `espresso`, `zenburn`, 
# `haddock`, `breezedark`, `textmate`, `arrow`, or `rstudio` or a file with extension `.theme`.
# Estimated time = 2h
title: "The 'vector' object."
output:
  learnr::tutorial:
    includes:
      before_body: !expr system.file(file.path("tutorials", "style.html"),package="rtrainer")
    theme: default
    highlight: default
    fig_caption: yes
    self_contained: true
    toc: yes
    toc_float: 
      toc_collapsed: false
    toc_depth: 4
    number_sections: false
    progressive: true
  html_document:
    theme: cosmo
    fig_caption: yes
    self_contained: yes
    toc: yes
    toc_float: 
      toc_collapsed: false
    toc_depth: 3
    number_section: true
  beamer_presentation:
    colortheme: dolphin
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    fonttheme: structurebold
    highlight: tango
    incremental: no
    keep_tex: no
    slide_level: 2
    theme: Montpellier
    toc: yes
  ioslides_presentation:
    colortheme: dolphin
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    fonttheme: structurebold
    highlight: tango
    incremental: no
    keep_md: no
    slide_level: 2
    smaller: yes
    theme: cerulean
    toc: yes
    widescreen: yes
  pdf_document:
    fig_caption: yes
    highlight: zenburn
    toc: yes
    toc_depth: 3
  slidy_presentation:
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    highlight: tango
    incremental: no
    keep_md: no
    self_contained: yes
    slide_level: 2
    smaller: yes
    theme: cerulean
    toc: yes
    widescreen: yes
font-import: http://fonts.googleapis.com/css?family=Risque
font-family: Garamond
transition: linear
runtime: shiny_prerendered
---

<!--  
Here the parameters about the documents.
https://www.rstudio.com/wp-content/uploads/2015/03/rmarkdown-reference.pdf 
--> 


```{css, echo=FALSE}

```

<script language="JavaScript" type="text/javascript">
          
          function sizeTbl2(h,i) {
          var tbl = document.getElementById("section-" + i);
          tbl.style.display = h;
          }

</script>



<style>
.exo {
  border-radius: 5px;
  margin-top: 5px;
  margin-bottom: 5px;
  padding-top: 5px;
  padding-bottom: 5px;
  padding-left: 5px;
  padding-right: 5px;
  background-color: #fcede3;
  color: rgb(51, 51,153);
}
.tips {
       padding-top: 5px;
       padding-bottom: 5px;
       padding-left: 5px;
       padding-right: 5px;
       border: 1px dashed #2f6fab;
       background-color: #EEFFEE;
}
.solution {
            margin-top: 5px;
            margin-bottom: 5px;
            padding-top: 5px;
            padding-bottom: 5px;
            padding-left: 5px;
            padding-right: 5px;
            border: 1px dashed #FFFFFF;
            background-color: #EEEEFF;
            color: #0000BB;
            font-size: 11px;
}
</style>


```{r setup, include=FALSE}
library(learnr)
library(gradethis)
knitr::opts_chunk$set(echo=TRUE, eval=TRUE, cache=FALSE, message=FALSE, warning=FALSE, comment="")
gradethis::gradethis_setup()
```

<!--  
<img src="img/graph_with_r.jpeg"> 
![](img/graph_with_r.jpeg)
--> 

## Modes

**All elements of a vector are of the same type**. In the R language, we speak of 'modes' (or modalities). There are four modes: **"numeric", "character", "logical" and "complex"**. 

### Numerical mode

We've already seen part of this. R lets you create numerical vectors. For example, below, we create a sequence of integers from 1 to 10, which is stored in a vector.

```{r vectornum0, exercise=TRUE}
integers <- 1:10
print(integers)
```

<div class="exo">

- Create a vector `z` containing the following integers in this order: 3, 12, 17, 20 to 41 and 50.

```{r integerexo, exercise=TRUE}

```

```{r integerexo-solution}
z <- c(3, 12, 17, 20:41, 50)
```

```{r integerexo-check }
gradethis::grade_result(
  pass_if(~ sum(z) == 753)
)
```
</div>


### Character mode

Below is an example of vectors containing *character strings*.

```{r vector0, exercise=TRUE}
names <- c("Stephen", "Katia", "Fairouz") 
print(names)
class(names) # <=> is(names) 
```

<div class="exo">

- Create `this_names`, a vector containing the following strings: “Riad”, “Robert” and “Nastassja”.

```{r charliste, exercise=TRUE}

```

```{r charliste-solution}
this_names <-  c("Riad", "Robert", "Nastassja") 
```

```{r charliste-check }
gradethis::grade_result(
  pass_if(~ this_names[1] == "Riad"),
  pass_if(~ this_names[2] == "Robert"),
  pass_if(~ this_names[3] == "Nastassja")
)
```

</div>

### Logic mode

Below is an example of vectors containing Booleans (variables with two states, 'true' and 'false' or 0 and 1). Note that true and false are represented by the keywords `TRUE` and `FALSE` (which are reserved terms in the language). 

```{r vector1, exercise=TRUE}
logic <- c(TRUE, FALSE, TRUE) 
print(logic)
class(logic) # is(logic)
```

`TRUE` and `FALSE` can be replaced by `T` and `F` respectively. However, this last solution is not recommended to increase readability.

```{r vector1_1, exercise=TRUE}
logic <- c(T, F, T)
print(logic)
class(logic) # is(logic)
```


<div class="exo">

- Create a `true_false` vector containing the Boolean values true, false, false, false, true.

```{r logique, exercise=TRUE}

```

```{r logique-solution}
true_false <- c(TRUE, FALSE, FALSE, FALSE, TRUE) 
```

```{r logique-check }
gradethis::grade_result(
  pass_if(~ sum(true_false) == 2),
  pass_if(~ length(true_false) == 5)
)
```

</div>

### Comparison operators

Logical vectors are generally created by R when comparing with comparison operators such as those in the table below:


| Operation | Result |                                                                  |
|-------|-------------------------------------------------------------------------------|
| a == b | Equality operator. Returns true if a and b contain equal values. |
| a != b	| Returns true if a and b contain different values. |
| a > b	| Returns true if the value of a is greater than the value of b. |
| a < b	| Returns true if the value of a is less than the value of b.  |
| a >= b	| Returns true if the value of a is greater than or equal to the value of b.  |
| a <= b	| Returns true if the value of a is less than or equal to the value of b. |
| !a	|  Logical negation (the opposite). |

A simple example is given below.

```{r bool1,  exercise=TRUE}
x <- c(1, 7, 10, 11, 3, -2, -4, 10)
x > 5
```

<div class="exo">

- Take the following vector x, test whether each of positions are less than 0 and put the result in a variable z. Print z.

```{r logicaltest, exercise=TRUE}
x <- c(2, -3, -2, 0, 6, 5)
```

```{r logicaltest-solution}
x <- c(2, -3, -2, 0, 6, 5)
z <- x < 0
print(z)
```

```{r logicaltest-check }
gradethis::grade_result(
  pass_if(~ all(z==c(FALSE, TRUE, TRUE, FALSE, FALSE, FALSE)))
)
```

</div>

You may also compare the paired positions of two vectors.

```{r bool1a,  exercise=TRUE, exercise.setup="bool1" }
x <- c(1, 7, 10, 11, 3, -2)
y <- c(2, 3, 10, 11, 4, -5)
x > y
```

<div class="exo">

- Given the following vectors `x` and `y`, store in `z` the Boolean vector indicating whether a position of `x` is equal (==) to a position of `y`.

```{r logicaltest_1, exercise=TRUE}
x <- c(2, -3, -2, 0, 6, 5)
y <- c(2, 3, 10, 11, 4, -5)
```

```{r logicaltest_1-solution}
x <- c(2, -3, -2, 0, 6, 5)
y <- c(2, 3, 10, 11, 4, -5)
z <- x == y
```

```{r logicaltest_1-check }
gradethis::grade_result(
  pass_if(~ all(z==c(TRUE, FALSE, FALSE, FALSE, FALSE, FALSE)))
)
```

</div>

<div class="exo">

- Given the following vectors `x` and `y`, store in z the Boolean vector indicating whether a position of `x` differs (!=) from that of `y`.

```{r logicaltest_2, exercise=TRUE}
x <- c(2, -3, -2, 0, 6, 5)
y <- c(2, 3, 10, 11, 4, -5)

```

```{r logicaltest_2-solution}
x <- c(2, -3, -2, 0, 6, 5)
y <- c(2, 3, 10, 11, 4, -5)
z <- x != y
```

```{r logicaltest_2-check }
gradethis::grade_result(
  pass_if(~ all(z==c(FALSE, TRUE, TRUE, TRUE, TRUE, TRUE)))
)
```

</div>

### Logic operators 

Logical vectors support logical operations, including the `&` operator (AND) and the `|` operator (OR). They can be used to test two vectors, **position by position**, to determine whether both positions are true (`&`) or whether at least one of them is true (`|`).

```{r vector1b, exercise=TRUE}
x <- c(TRUE, FALSE, TRUE, FALSE)
y <- c(TRUE, FALSE, FALSE, TRUE)
print(x & y)
print(x | y)
```

<div class="exo">
- Store in z the logic vector corresponding to the test: x > 1 and y < 7

```{r logicaltest_3, exercise=TRUE}
x <- c(1, 7, 10, 11, 3, -2, -4, 10)
y <- c(2, 3, 10, 11, 4, -5, -4, 12)
```

```{r logicaltest_3-solution}
x <- c(1, 7, 10, 11, 3, -2, -4, 10)
y <- c(2, 3, 10, 11, 4, -5, -4, 12)
z <- x > 1 & y < 7
print(z)
```

```{r logicaltest_3-check}
gradethis::grade_result(
  pass_if(~identical(z, c(FALSE, TRUE, FALSE, FALSE, TRUE, FALSE, FALSE, FALSE)))
)
```
</div>

### Any() and all() functions

With these function we can test whether at least one (**any()**) or all (**all()**) of the positions of a Boolean vector are true.

```{r exo1, exercise=TRUE}
x <- 1:10
x > 0
all(x > 0)
```

```{r exob, exercise=TRUE}
x <- -10:10
x > 0
all(x > 0)
any(x > 0)
```

### The which() function

The `which()` function returns **the positions** of a vector for which a condition is `TRUE`.

```{r exo2, exercise=TRUE}
x <- c(2, -3, 5, 7, 10, 2, -4, 11)
which(x < 0)
```

<div class="exo">
- Store in `z` all positions for which the values of `x` and `y` are the same

```{r vector1c, exercise=TRUE}
x <- c(1, 7, 10, 11, 3, -2, -4, 10)
y <- c(2, 3, 10, 11, 4, -5, -4, 12)
```

```{r vector1c-solution}
x <- c(1, 7, 10, 11, 3, -2, -4, 10)
y <- c(2, 3, 10, 11, 4, -5, -4, 12)
z <- which(x == y)
```

```{r vector1c-check}
gradethis::grade_result(
  pass_if(~ sum(z) == 14),
  pass_if(~ length(z) == 3)
)
```
</div>

<div class="exo">
- Store in vector `z` all positions for which the values of `x` and `y` differ.

```{r vector1c2, exercise=TRUE}
x <- c(1, 7, 10, 11, 3, -2, -4, 10)
y <- c(2, 3, 10, 11, 4, -5, -4, 12)
```

```{r vector1c2-solution}
x <- c(1, 7, 10, 11, 3, -2, -4, 10)
y <- c(2, 3, 10, 11, 4, -5, -4, 12)
z <- which(x != y)
```

```{r vector1c2-check}
gradethis::grade_result_strict(
  pass_if(~ sum(z) == 22),
  pass_if(~ length(z) == 5)
)
```
</div>

<div class="exo">
- Store in vector `z` all positions for which `y` values are at least two times greater than `x`.

**NB**: The multiplication operator in R is '*'.

```{r vector1c3, exercise=TRUE}
x <- c(1, 7, 10, 11, 3, -2, -4, 10)
y <- c(2, 3, 10, 11, 4, -5, -4, 12)
```

```{r vector1c3-solution}
x <- c(1, 7, 10, 11, 3, -2, -4, 10)
y <- c(2, 3, 10, 11, 4, -5, -4, 12)
z <- which(y >= 2 * x)
```

```{r vector1c3-check}
gradethis::grade_result_strict(
  pass_if(~ sum(z) == 8), 
  pass_if(~ length(z) == 2)
  
)
```
</div>

## Example of functions to create vectors.

### rep()

There are numerous functions for creating vectors. For example, the `rep()` function repeats an existing vector or value.

```{r repeat_1, exercise=TRUE}
rep(x=1:5, times=3)
```

<div class="exo">
Use the `rep()` (**rep**eat) and `c()` functions to store the following vector in the `dna` variable:

```{r vector2, eval=T, echo=F}
print(rep(c("A", "T", "T"), 10))
```

```{r vector3, exercise=TRUE}
x <- c("A", "T", "T")
```

```{r vector3-solution}
x <- c("A", "T", "T")
dna <- rep(x, 10)
```

```{r vector3-check}
gradethis::grade_result(
  pass_if(~identical(dna, c("A", "T", "T", "A", "T", "T", "A", "T", "T", "A", "T", "T", 
"A", "T", "T", "A", "T", "T", "A", "T", "T", "A", "T", "T", "A", 
"T", "T", "A", "T", "T")))
)
```
</div>

### The seq() function

The `seq()` function is used to create sequences of numerical values at regular intervals. You define a starting value (`from`), an ending value (`to`) and increment according to a defined step (`by`) or a targeted number of values (`length.out`).

<div class="exo">

- Look at the help on the `seq()` function and propose :

  - An instruction to create a vector `a` containing regularly spaced values (in steps of 0.1) between 1 and 10 (argument `by`).
  
  - An instruction to create a vector `b` containing 20 values regularly spaced between 1 and 10 (argument `length.out`).  

```{r seq1, exercise=TRUE}

```

```{r seq1-solution}
a <- seq(from=1, to=10, by=0.1)
print(a)
b <- seq(from=1, to=10, length.out=20)
print(b)
```

```{r seq1-check}
gradethis::grade_result(
  pass_if(~ sum(a) == 505),
  pass_if(~ round(sum(b), 2) == 110)
)
```
</div>


## Generating random values 

### Functions for generating random values

R allows you to generate random numbers (*i.e.* **r**andom) from a wide variety of distributions (*e.g.* normal distribution, uniform distribution, Poisson distribution...). 

The code below randomly draws 1000 values from a normal distribution with mean 0 and standard deviation 1. 

**NB:** In the computer context, the values appear random, but are in fact calculated from algorithms that attempt to reproduce random processes. The "seed" of the `set.seed()` function (`123` in the example below) allows you to force the algorithm to return the same "random" values (as if you could somehow force randomness...). 


```{r vector4, exercise=TRUE}
set.seed(123)
x <- round(rnorm(n=1000, mean=0, sd=1), digit=2) # use print(x) to show x
length(x)
```

You can use `head()` to show the first 6 elements of x

```{r vector4_1, exercise=TRUE, exercise.setup="vector4"}
head(x)
```

<div class="exo">
- Store in `z` 500 values drawn from a normal distribution with a mean of 2 and a standard deviation of 3. The values should be rounded (`round()`) to 2 decimal places (use `round()`).

```{r rnorm1, exercise=TRUE}
set.seed(123) # Keep this to ensure everyone has the same results

```

```{r rnorm1-solution}
set.seed(123)
z <- round(rnorm(500, mean = 2, sd=3), 2)
```

```{r rnorm1-check}
gradethis::grade_result(
  pass_if(~ sum(z) > 1000),
  pass_if(~ length(z) == 500)
)
```
</div>

### Visualizing the distributions

You can plot the histogram of a distribution with the `hist()` function. The `breaks` argument controls the **number of intervals** in the histogram.

```{r vector5, exercise=TRUE, exercise.setup = "vector4"}
x <- round(rnorm(n=1000, mean=0, sd=1), digit=2)
hist(x, breaks = 20, col="blue", border="white")
```

<div class="exo">
Plot the distribution of `x` with `hist()` using 5 and 100 for the breaks parameter. What do you observe?

```{r vector5b, exercise=TRUE, exercise.setup = "vector4"}

```

```{r vector5b-solution}
hist(x, breaks = 5)
hist(x, breaks = 100)
```
</div>


## Sorting operations

### The sort() function

The `sort()` function allows you to sort a vector with the **decreasing** argument indicating whether to perform an **ascending** or **descending** sorting. 

```{r vector7, exercise=TRUE}
x <- c(20:25, 1:5, 50)
sort(x, decreasing = TRUE)
```

<div class="exo">
Store in `x` the values of `x` ordered in ascending order.

```{r vector8, exercise=TRUE, exercise.setup ="vector7"}

```

```{r vector8-solution}
x <- sort(x, decreasing = FALSE)
```

```{r vector8-check }
gradethis::grade_result(
  pass_if(~identical(x, c(1, 2, 3, 4, 5, 20, 21, 22, 23, 24, 25, 50)))
)
```
</div>

## Random sampling from a vector

### The sample() function

The `sample()` function randomly selects `size` elements from a vector `x`. With `replace = TRUE`, each drawn element is replaced in `x` before the next draw.

```{r vector9, exercise=TRUE}
sample(x=1:10, size = 5, replace = TRUE)
```

### Exercises around the sample() function

<div class="exo">
- Perform 2 random selection of 5 elements with replacement from a vector containing all integers from 1 to 10.

```{r vector10, exercise=TRUE}

```

```{r vector10-solution}
print(sample(x=1:10, size = 5, replace = TRUE))
print(sample(x=1:10, size = 5, replace = TRUE))
```
</div>

```{r vector10q, echo=FALSE }
learnr::question("Do the returned values change from one call to another?",
  answer("No"),
  answer("Yes", correct = TRUE),
  allow_retry = TRUE 
)
```

```{r vector10a2, echo=FALSE }
learnr::question("In the generated vectors, could you observe any duplicate values?",
  answer("No"),
  answer("Yes", correct = TRUE),
  allow_retry = TRUE 
)
```

<div class="exo">
- Perform 2 random draws of 5 elements **without replacement** from a vector containing all integers from 1 to 10.

```{r vector11, exercise=TRUE}

```

```{r vector11-solution}
print(sample(x=1:10, size = 5, replace = FALSE))
print(sample(x=1:10, size = 5, replace = FALSE))
```

</div>

```{r vector11q, echo=FALSE }
learnr::question("In each of the generated vectors, could you observe any identical values?",
  answer("Yes"),
  answer("No", correct = TRUE),
  allow_retry = TRUE 
)
```

## Mathematical operations

### Mathematical operators

**Mathematical operations** on vectors use standard operators (+, -, \*, /, ^ for raising to a power) and are **vectorized**, meaning they apply to all elements automatically, eliminating the need for explicit loops. This simplifies syntax and enables complex operations in just a few lines.

To add 10 to all elements of a numeric vector, you would thus write:

```{r vector22, exercise=TRUE, eval=TRUE}
x <- 1:10
print(x + 10)
```

To sum the elements at each position *p* of two vectors `x` and `y`, you would write:

```{r vector22b, exercise=TRUE, eval=TRUE}
x <- 1:10
y <- 11:20
print(x + y)
```

The same principle will apply to other operations:

```{r vector23, exercise=TRUE, eval=FALSE, exercise.setup="vector22b"}
x ^ 2
x * 3
x / 4
x * y
```

<div class="exo">

- Compute the sum of each position of `x` and `y` and store the result in `z`.

```{r exovectorisation, exercise=TRUE }
x <- 1:5
y <- 5:1
```

```{r exovectorisation-solution}
x <- 1:5
y <- 5:1
z <- x + y
```

```{r exovectorisation-check }
gradethis::grade_result(
  pass_if(~ sum(z)==30)
)
```
</div>

<div class="exo">

- Consider the following example:

```{r notworking, exercise=TRUE }
x <- 1:3
y <- 1:8
x + y
```

```{r qnw1, echo=FALSE }
learnr::question("What happens in the previous example?",
  answer("R refuses to add two vectors of different sizes because the size of one is not a multiple of the other."),
  answer("R accepts (with a warning) to add two vectors of different sizes when the size of one is not a multiple of the other.", correct = TRUE),
  allow_retry = TRUE 
)
```

</div>

<div class="exo">

- Given the following example:

```{r notworking_2, exercise=TRUE }
x <- 1:4
y <- 1:8
x + y
```

```{r qnw2, echo=FALSE }
learnr::question("What happens in the previous example?",
  answer("R accepts the addition operation without warning because the size of one is a multiple of the size of the other.", correct = TRUE),
    answer("R refuses to add two vectors of different sizes because their sizes differ."),
  allow_retry = TRUE 
)
```

</div>

## Missing values

A vector can contain missing values. In R, they are denoted as `NA` (**Not Available**).

```{r vector23c, exercise=TRUE, eval=TRUE}
x <- c(NA, 40, 2, NA, 7, 8, 30)
print(x)
```

Other undefined values exist such as `Inf` (**Infinite**), `-Inf`, or `NaN` (**NotANumber**). They are generally produced by calculation errors.

```{r vector23d, exercise=TRUE, eval=TRUE}
1/0
log(0)
0/0
```

The functions `is.na()`, `is.infinite()`, and `is.nan()` allow you to iterate over vectors and return a boolean value for each position indicating whether the position is of type `NA`, `Inf/-Inf`, or `NaN`, respectively.

```{r vector23db, exercise=TRUE, eval=TRUE}
is.na(c(1:3, NA, 1:3))
is.nan(c(1:3, NaN, 1:3))
is.infinite(c(1:3, Inf, 1:3))
```


<div class="exo">
- Given the vector below, store in `z` the number of position that are `NA` (use `is.na()`).

**Tips**: You can apply `sum()` over a boolean vector to count the TRUE positions.

```{r vector23e, exercise=TRUE}
set.seed(123)
x <- sample(c(1:10, NA, -Inf, Inf), size = 1000, replace = TRUE)
```

```{r vector23e-solution }
set.seed(123)
x <- sample(c(1:10, NA, -Inf, Inf), size = 1000, replace = TRUE)
z <- sum(is.na(x))
```

```{r vector23e-check }
gradethis::grade_result(
  pass_if(~ z == 70)
)
```
</div>


## Quizz

Answer the following questions:

<div class="exo">
```{r exoq1, echo=FALSE }
learnr::question("How many modes exist in R?",
  answer("0"),
  answer("2"),
  answer("4", correct = TRUE),
  answer("10"),
  allow_retry = TRUE 
)
```

```{r exoq2, echo=FALSE }
learnr::question("Select the logical values accepted by the R language.",
  answer("T",  correct = TRUE),
  answer("true"),
  answer("t"),
  answer("F",  correct = TRUE),
  answer("false"),
  answer("f"),
  answer("TRUE", correct = TRUE),
  answer("FALSE", correct = TRUE),
  allow_retry = TRUE 
)
```

```{r exoq3, echo=FALSE }
learnr::question("Select the items that are logical operators in R.",
  answer("="),
  answer("\\>=",  correct = TRUE),
  answer("==",  correct = TRUE),
  answer("<=", correct = TRUE),
  answer("!=!"),
  answer("!=", correct = TRUE),
  answer("==!"),
  answer("<", correct = TRUE),
  answer("\\>", correct = TRUE),
  allow_retry = TRUE 
)
```

```{r exoq4, echo=FALSE }
learnr::question("If x and y are two numerical vectors, which instructions may be accepted by R.",
  answer("x > y",  correct = TRUE),
  answer("x * y > 0",  correct = TRUE),
  answer("x / y < 10",  correct = TRUE),
  answer("8 < x < 10"),
    answer("x & y",  correct = TRUE),
  answer("x^2 + y^2 < 1",  correct = TRUE),
  answer("x > 10 & y < 3",  correct = TRUE),
  answer("!x > 3 & y < 10", correct = TRUE),
  allow_retry = TRUE 
)
```

</div>


### The end

Thank you for following this tutorial.

<!--```{r exo10, eval=TRUE, echo=FALSE } 
htmltools::tags$button("Retour au Menu", 
                       class="btn btn-primary",
                       onclick = "window.location.href='https://w3docs.com';"
                        )
```
-->




