---
# theme
# "bootstrap", "cerulean", "cosmo", "darkly", "flatly", "journal", "lumen", "paper", 
# "readable", "sandstone", "simplex", "spacelab", "united", "yeti"
# highlight: `default`, `tango`, `pygments`, `kate`, `monochrome`, `espresso`, `zenburn`, 
# `haddock`, `breezedark`, `textmate`, `arrow`, or `rstudio` or a file with extension `.theme`.

title: "The 'factor' object"
output:
  learnr::tutorial:
    includes:
      before_body: !expr system.file(file.path("tutorials", "style.html"),package="rtrainer")
    theme: default
    highlight: default
    fig_caption: yes
    self_contained: true
    toc: yes
    toc_float: 
      toc_collapsed: false
    toc_depth: 4
    number_sections: false
    progressive: true
  html_document:
    theme: cosmo
    fig_caption: yes
    self_contained: yes
    toc: yes
    toc_float: 
      toc_collapsed: false
    toc_depth: 3
    number_section: true
  beamer_presentation:
    colortheme: dolphin
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    fonttheme: structurebold
    highlight: tango
    incremental: no
    keep_tex: no
    slide_level: 2
    theme: Montpellier
    toc: yes
  ioslides_presentation:
    colortheme: dolphin
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    fonttheme: structurebold
    highlight: tango
    incremental: no
    keep_md: no
    slide_level: 2
    smaller: yes
    theme: cerulean
    toc: yes
    widescreen: yes
  pdf_document:
    fig_caption: yes
    highlight: zenburn
    toc: yes
    toc_depth: 3
  slidy_presentation:
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    highlight: tango
    incremental: no
    keep_md: no
    self_contained: yes
    slide_level: 2
    smaller: yes
    theme: cerulean
    toc: yes
    widescreen: yes
font-import: http://fonts.googleapis.com/css?family=Risque
font-family: Garamond
transition: linear
runtime: shiny_prerendered
---

<!--  
Here the parameters about the documents.
https://www.rstudio.com/wp-content/uploads/2015/03/rmarkdown-reference.pdf 
--> 


```{css, echo=FALSE}

```

<script language="JavaScript" type="text/javascript">
          
          function sizeTbl2(h,i) {
          var tbl = document.getElementById("section-" + i);
          tbl.style.display = h;
          }

</script>



<style>
.exo {
  border-radius: 5px;
  margin-top: 5px;
  margin-bottom: 5px;
  padding-top: 5px;
  padding-bottom: 5px;
  padding-left: 5px;
  padding-right: 5px;
  background-color: #fcede3;
  color: rgb(51, 51,153);
}
.tips {
       padding-top: 5px;
       padding-bottom: 5px;
       padding-left: 5px;
       padding-right: 5px;
       border: 1px dashed #2f6fab;
       background-color: #EEFFEE;
}
.solution {
            margin-top: 5px;
            margin-bottom: 5px;
            padding-top: 5px;
            padding-bottom: 5px;
            padding-left: 5px;
            padding-right: 5px;
            border: 1px dashed #FFFFFF;
            background-color: #EEEEFF;
            color: #0000BB;
            font-size: 11px;
}
</style>


```{r setup, include=FALSE}
library(learnr)
library(gradethis)
knitr::opts_chunk$set(echo=TRUE, eval=TRUE, cache=FALSE, message=FALSE, warning=FALSE, comment="")
gradethis::gradethis_setup()
```

<!--  
<img src="img/graph_with_r.jpeg"> 
![](img/graph_with_r.jpeg)
--> 


## Factors

In R, **factors** represent **categorical variables** created from character strings or numerical vectors. They are essential for tasks like:

- Compute statistics (*e.g* mean, variance, t-test...) according to categories
- Create diagrams (*e.g* distributions, boxplot...) according to categories
- ...

As an example, you may be interested to analyze values according to any membership. For example, below, we create an ordered vector containing the treatments received by 60 patients. 

```{r fact1, exercise=TRUE}
# ttmt is, for the moment, a simple vector of character strings.
ttmt <- sort(sample(c("Placebo", "hydroxychloroquine", "dexamethasone"), size=60, replace = TRUE))
class(ttmt)
print(head(ttmt))
```

By default, this vector contains no information about any categories.

### The as.factor() function

The `as.factor()` function converts a vector into a factor (*i.e* categorical variable).

```{r exo2, exercise=TRUE, exercise.setup = "fact1" }
# ttmt is tranformed/casted into a factor
ttmt <- as.factor(ttmt)
class(ttmt)
```


```{r quizz_fact_1, echo=FALSE }
learnr::question("Let's print the ttmt factor. How does R indicate that it's a factor ?",
  answer("It indicates ‘Levels: dexamethasone hydroxychloroquine Placebo’ on the first line", correct = FALSE),
  answer("It indicates ‘Levels: dexamethasone hydroxychloroquine Placebo’ on the last line", correct = TRUE ),
  allow_retry = TRUE
)
```

### The levels() function

The `levels()` function lets you extract the names of the categories (the **levels/modalities** of these categories).

```{r exo3, exercise=TRUE, exercise.setup="exo2" }
levels(ttmt)
```

You can change the category names as shown in the following example. Check, using `print()`, that the levels have changed.

```{r exo3b, exercise=TRUE, exercise.setup="exo2"}
# We can change the category names.
levels(ttmt) <- c("dex", "hcq", "cont")
levels(ttmt)
```

### The table() function

The `table()` function returns the size of each category. For instance, let say we have information related to patients genders.

```{r exo4, exercise=TRUE, exercise.setup="exo3b" }
set.seed(123)
gender <- sample(c("Male", "Female"), size = 60, replace = T)
gender <- as.factor(gender)
head(gender)
```

The `table()` function can be used to count the number of men and women.

```{r exo5, exercise=TRUE, exercise.setup="exo4" }
table(gender)
```

Now that we have two variables (`ttmt`and `gender`) we can also easily cross the variables and create a contengency matrix to get the number of men/women expose to the various treatments.


```{r exo6a, exercise=TRUE, exercise.setup="exo5" }
table(ttmt, gender)
```

### Exercise with the table() function

<div class="exo">

- Given the `age` factor (see below), count the number of individuals in the different levels through the factors `ttmt`, `gender` and `age`.

**NB:** The first argument to table (*cf* `help(table)`) is '...', which means you can pass as many factors as you like as arguments.

```{r exo5c, exercise=TRUE , exercise.setup="exo5"}
set.seed(456)
age <- sample(c("<10", "10-30", "30-50", ">50"), size = 60, replace = T)
```

```{r exo5c-solution}
set.seed(456)
age <- sample(c("<10", "10-30", "30-50", ">50"), size = 60, replace = T)
table(ttmt,  gender, age)
```

</div>


## Ordinal variables

An **ordinal variable** is similar to a **categorical variable but has a defined order** (e.g., education level: "high school," "BS," "MS," "PhD"; or rating: "like," "neutral," "dislike").

### Creating ordinal variables

For example, `age` could be naturally converted to an ordinal variable using the `factor()` function (which contains `ordered` arguments compared to `as.factor()`). 

```{r exo6, exercise=TRUE, exercise.setup="exo5c" }
age <- factor(age, levels = c("<10", "10-30", "30-50", ">50"), ordered = TRUE)
age
```

```{r exo6b1, echo=FALSE }
learnr::question("When printing the age object, how does R indicate that it's an ordinal  variable?",
  answer("It indicates 'Levels: <10 < 10-30 < 30-50 < >50' on the first line", correct = FALSE),
  answer("It says 'Levels: <10 < 10-30 < 30-50 < >50' on the last line ?", correct = TRUE),
  allow_retry = TRUE 
)
```

```{r exo6b0, echo=FALSE }
learnr::question("What does the '<' in 'Levels: cont < hcq < dex' stand for?",
  answer("It indicates that the modalities can be classified in a certain natural order to enable comparison.", correct = TRUE),
  answer("It allows to separate doesn't make sense."),
  answer("It indicates that the modalities cannot be sorted."),
  allow_retry = TRUE 
)
```

## Applying operations based on catagories

### The tapply() function

Factors enable operations on vectors by category. For example, `tapply()` applies functions like `mean()` to each category. Below, it calculates the mean and visualizes infection distributions across three patient treatment groups.

```{r exo7, exercise=TRUE, exercise.setup="exo6" } 
# infection is a vector containing the infection levels
# infection levels of each patient
infection <- round(c(rnorm(20, mean = 3, sd=1), 
                      rnorm(20, mean = 10, sd=2), 
                      rnorm(20, mean = 10, sd=1)), 
                    2)
tapply(infection, ttmt, mean)
```

This mean can also be calculated as a function of several categorical variables.

```{r exo8, exercise=TRUE, exercise.setup="exo7" } 
tapply(infection, list(ttmt, gender), mean)
```

These factors are also extremely useful for graphical operations. Below, for example, we've created boxplot based on patient categories using very basic R functions. The `~` (tilda) can be read as 'as a function of' or 'according to'.

```{r exo9, exercise=TRUE, exercise.setup="exo7" } 
boxplot(infection ~ ttmt)
```

## Creating factors with the cut() function

The `cut()` function splits a numeric vector into intervals, encoding values based on their intervals, effectively converting a numeric variable into categories.

```{r exo10_0, exercise=TRUE, exercise.setup="exo7" }
infection_class <- cut(infection, breaks = 4)
is(infection_class)
print(levels(infection_class))
```

Now that we've created infection level classes, we could create the corresponding contingency table:

```{r exo10_1, exercise=TRUE, exercise.setup="exo10_0"}
cont_table <- table(infection_class, ttmt)
cont_table
```

We'll look at the use of this *factor* object in more detail later.

## Exercises

### Exercise 1

<div class="exo">

- Given the object below.

```{r exo10, exercise=TRUE }
chromosome <- c(11, 2, 7, 7, 8, 10, 2, 20, 22, 1, 3, 10, 10, 11, 20)
```

```{r exo11, echo=FALSE}
learnr::question("What is the type of the chromosome object?",
  answer("factor"),
  answer("character"),
  answer("numeric", correct = TRUE),
  answer("boolean"),
  allow_retry = TRUE
)
```

```{r exo12, echo=FALSE}
learnr::question("What are the associated categories (levels()) ?",
  answer("Zero"),
  answer("Several"),
  answer("NULL", correct = TRUE),
  allow_retry = TRUE
)
```

- Transform `chromosome` with the following code.

```{r exo13, exercise=TRUE, exercise.setup="exo10" } 
chromosome <- as.factor(chromosome)
```

```{r exo14, echo=FALSE } 
learnr::question("What is the type of chromosome ?",
  answer("factor", correct = TRUE),
  answer("character"),
  answer("numeric"),
  answer("boolean"),
  allow_retry = TRUE 
)
```

```{r exo15, echo=FALSE }
learnr::question("How many levels does it contains?",
  answer("0", correct = FALSE),
  answer("9", correct = TRUE),
  answer("10", correct = FALSE),  
  answer("NULL", correct = FALSE),
  allow_retry = TRUE
)
```

</div>

### Exercise 2

<div class="exo">

- Given the *codons* object below. Complete the code to print the number of "ATG" codons. Store this value in a variable named *att_nb*.

```{r exo16, exercise=TRUE}
set.seed(123)
codons <- sample(c("ATT", "ATG", "TTG", "TCG", "GCG", "TGA"), size=1000, replace = TRUE)
codons <- as.factor(codons)
```


```{r exo16-solution }
set.seed(123)
codons <- sample(c("ATT", "ATG", "TTG", "TCG", "GCG", "TGA"), size=1000, replace = TRUE)
codons <- as.factor(codons)
tb_codon <- table(codons)
att_nb <- tb_codon['ATG']
```

```{r exo16-check }
gradethis::grade_result(
  pass_if(~ setNames(att_nb, NULL) ==  176)
)
```
</div>

<div class="exo">

- Propose an instruction to transform the codon categories "ATG" into "START" and "TGA" into "STOP".

```{r exo17,  exercise=TRUE, echo=FALSE, exercise.setup="exo16" } 
set.seed(123)
codons <- sample(c("ATT", "ATG", "TTG", "TCG", "GCG", "TGA"), size=1000, replace = TRUE)
codons <- as.factor(codons)
```


```{r exo17-solution}
# Look at the order 
levels(codons)
# Change the names of the levels
levels(codons) <- c("START", "ATT", "GCG", "TCG", "STOP", "TTG")
```


```{r exo17-check }
gradethis::grade_result(
  pass_if(~  table(codons)["START"] == 176),
  pass_if(~  table(codons)["STOP"] == 164)
)
```
</div>


### End of the section

Thank you for following this tutorial.


<!--
  $$```{r end, eval=TRUE, echo=FALSE } 
 htmltools::tags$button("Retour au Menu", 
                        class="btn btn-primary",
                        onclick = "window.location.href='https://w3docs.com';"
                         )
 $$```
 -->

