---
# theme
# “bootstrap”, “cerulean”, “cosmo”, “darkly”, “flatly”, “journal”, “lumen”, “paper”, 
# “readable”, “sandstone”, “simplex”, “spacelab”, “united”, “yeti”
# highlight: `default`, `tango`, `pygments`, `kate`, `monochrome`, `espresso`, `zenburn`, 
# `haddock`, `breezedark`, `textmate`, `arrow`, or `rstudio` or a file with extension `.theme`.

title: "Lists"
author: "D. Puthier"
output:
  learnr::tutorial:
    includes:
      before_body: !expr system.file(file.path("tutorials", "style.html"),package="rtrainer")
    theme: default
    highlight: default
    fig_caption: yes
    self_contained: true
    toc: yes
    toc_float: 
      toc_collapsed: false
    toc_depth: 4
    number_sections: false
    progressive: false
  html_document:
    theme: cosmo
    fig_caption: yes
    self_contained: yes
    toc: yes
    toc_float: 
      toc_collapsed: false
    toc_depth: 3
    number_section: true
  beamer_presentation:
    colortheme: dolphin
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    fonttheme: structurebold
    highlight: tango
    incremental: no
    keep_tex: no
    slide_level: 2
    theme: Montpellier
    toc: yes
  ioslides_presentation:
    colortheme: dolphin
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    fonttheme: structurebold
    highlight: tango
    incremental: no
    keep_md: no
    slide_level: 2
    smaller: yes
    theme: cerulean
    toc: yes
    widescreen: yes
  pdf_document:
    fig_caption: yes
    highlight: zenburn
    toc: yes
    toc_depth: 3
  slidy_presentation:
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    highlight: tango
    incremental: no
    keep_md: no
    self_contained: yes
    slide_level: 2
    smaller: yes
    theme: cerulean
    toc: yes
    widescreen: yes
font-import: http://fonts.googleapis.com/css?family=Risque
font-family: Garamond
transition: linear
runtime: shiny_prerendered
---

<!--  
Here the parameters about the documents.
https://www.rstudio.com/wp-content/uploads/2015/03/rmarkdown-reference.pdf 
--> 


```{css, echo=FALSE}

```

<script language="JavaScript" type="text/javascript">
          
          function sizeTbl2(h,i) {
          var tbl = document.getElementById("section-" + i);
          tbl.style.display = h;
          }

</script>



<style>
.exo {
  border-radius: 5px;
  margin-top: 5px;
  margin-bottom: 5px;
  padding-top: 5px;
  padding-bottom: 5px;
  padding-left: 5px;
  padding-right: 5px;
  background-color: #fcede3;
  color: rgb(51, 51,153);
}
.tips {
       padding-top: 5px;
       padding-bottom: 5px;
       padding-left: 5px;
       padding-right: 5px;
       border: 1px dashed #2f6fab;
       background-color: #EEFFEE;
}
.solution {
            margin-top: 5px;
            margin-bottom: 5px;
            padding-top: 5px;
            padding-bottom: 5px;
            padding-left: 5px;
            padding-right: 5px;
            border: 1px dashed #FFFFFF;
            background-color: #EEEEFF;
            color: #0000BB;
            font-size: 11px;
}
</style>

```{r echo=FALSE}
# chunk below enables printing whole tutorial from browser e.g. to pdf
# DO NOT put any #comments in the chunk below, that stops it from working !! 
# from https://github.com/rstudio/learnr/issues/465
# saving csss in a separate file print.css didn't work locally or on shinyapps because browser couldn't find file 
```

```{css echo=FALSE}
@media print {
  .topicsContainer,
  .topicActions,
  .exerciseActions .skip {
    display: none;
  }
  .topics .tutorialTitle,
  .topics .section.level2,
  .topics .section.level3:not(.hide) {
    display: block;
  }
  .topics {
    width: 100%;
  }
  .tutorial-exercise, .tutorial-question {
    page-break-inside: avoid;
  }
  .section.level3.done h3 {
    padding-left: 0;
    background-image: none;
  }
  .topics .showSkip .exerciseActions::before {
    content: "Topic not yet completed...";
    font-style: italic;
  }
}  
  
```


```{r setup, include=FALSE}
library(learnr)
library(gradethis)
knitr::opts_chunk$set(echo=TRUE, eval=TRUE, cache=FALSE, message=FALSE, warning=FALSE, comment="")
gradethis::gradethis_setup()
```


<!--  
<img src="img/graph_with_r.jpeg"> 
![](img/graph_with_r.jpeg)
--> 

## Lists

### What are Lists?

The `list` object is **the most flexible** of R's basic objects. It has **no size constraints**, meaning you can store **objects of different sizes** within it. Similarly, it has **no type constraints**, allowing you to **store a wide variety of objects**. Each element in the list can have a name. These names can be accessed using the `names()` function, just like with vectors.


```{r exo1, exercise=TRUE}
my_list <- list(a_vector = 1:5, 
                a_matrix = matrix(8:13, nr = 2), 
                some_letters = LETTERS[10:15])
my_list
```

```{r exo0_1, exercise=TRUE, exercise.setup="exo1" }
names(my_list)
```

The size of the list (the number of stored objects) is given, as with vectors, by the `length()` function.

```{r exo0, exercise=TRUE, exercise.setup="exo1" }
length(my_list)
```

## Indexing Lists

### The Indexing Function '['

Note that in the case of lists, the indexing function ("[") returns a sublist. It is therefore useful for creating a new list that is a subset of the original list.


```{r exo2, exercise=TRUE, exercise.setup="exo1"}
sub_list <- my_list[c(1,3)]
is(sub_list)
sub_list
```


<!-- ```{r exo2_1, exercise=TRUE, exercise.setup="exo1" } -->
<!-- sub_list <- my_list[1] -->
<!-- is(sub_list) -->
<!-- sub_list -->
<!-- ``` -->

### The Indexing Function '[['

To directly access the elements contained in the list, you need to use the "[[" indexing function and pass it either a numeric value or the name of a target element as an argument. This function only accepts an atomic vector (containing a single element) as input.

```{r exo_3, exercise=TRUE, exercise.setup="exo1" }
# Mêmes résultats
my_list[[1]]
my_list[['a_vector']]
```


### The Indexing Function $

As with `data.frame` objects, you can also use the '$' operator to index the list.

```{r exo_4, exercise=TRUE, exercise.setup="exo1"  }
# Mêmes résultats
my_list$a_vector
```

<div class="exo">

- Using the indexing functions, extract the $5^{th}$ position of the vector `a_vector` contained in `my_list`. Store the result in the variable `a`.
- Using the indexing functions, extract the cell at row 2 / column 2 of the matrix `a_matrix`. Store the result in the variable `b`.

```{r exolist1, exercise=TRUE }
my_list <- list(a_vector = 1:5, 
                a_matrix = matrix(8:13, nr = 2), 
                some_letters = LETTERS[10:15])

```

```{r exolist1-solution}
my_list <- list(a_vector = 1:5, 
                a_matrix = matrix(8:13, nr = 2), 
                some_letters = LETTERS[10:15])


# Exo a: solution 1
a <- my_list[['a_vector']][5]
# Exo a: solution 2
v <- my_list[['a_vector']]
a <- v[5]
# Exo a: solution 3
v <- my_list$a_vector
a <- v[5]
# Exo a: solution 4
a <- my_list$a_vector[5]


# Exo b (your turn to find 
# additional solutions)
b <- my_list$a_matrix[2, 2]
```

```{r exolist1-check }
gradethis::grade_result(
  pass_if(~ a == 5),
  pass_if(~ b == 11)
)
```
</div>


<div class="exo">

- Using the indexing functions, extract the $5^{th}$ position of the vector `a_vector` contained in the nested list. Store the result in the variable `a`.
- Using the indexing functions, extract the cell at row 2 / column 2 of the matrix `a_matrix` contained in the nested list. Store the result in the variable `b`.

```{r exolist2, exercise=TRUE }
my_list <- list(A=list(a_vector = 1:5, 
                       a_matrix = matrix(8:13, nr = 2)),
                B=list(some_letters = LETTERS[10:15],
                       other_letters = LETTERS[18:26])
               )
```

```{r exolist2-solution}
my_list <- list(A=list(a_vector = 1:5, 
                       a_matrix = matrix(8:13, nr = 2)),
                B=list(some_letters = LETTERS[10:15],
                       other_letters = LETTERS[18:26])
               )

a <- my_list$A$a_vector[5]
b <- my_list$A$a_matrix[2, 2]
```

```{r exolist2-check }
gradethis::grade_result(
  pass_if(~ a == 5),
  pass_if(~ b == 11)
)
```
</div>

### End of the section

Thank you for following this tutorial.



