---
# theme
# “bootstrap”, “cerulean”, “cosmo”, “darkly”, “flatly”, “journal”, “lumen”, “paper”, 
# “readable”, “sandstone”, “simplex”, “spacelab”, “united”, “yeti”
# highlight: `default`, `tango`, `pygments`, `kate`, `monochrome`, `espresso`, `zenburn`, 
# `haddock`, `breezedark`, `textmate`, `arrow`, or `rstudio` or a file with extension `.theme`.

title: "Indexation des vecteurs"
output:
  learnr::tutorial:
    theme: default
    highlight: default
    fig_caption: yes
    self_contained: true
    toc: yes
    toc_float: 
      toc_collapsed: false
    toc_depth: 4
    number_sections: false
    progressive: true
  html_document:
    theme: cosmo
    fig_caption: yes
    self_contained: yes
    toc: yes
    toc_float: 
      toc_collapsed: false
    toc_depth: 3
    number_section: true
  beamer_presentation:
    colortheme: dolphin
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    fonttheme: structurebold
    highlight: tango
    incremental: no
    keep_tex: no
    slide_level: 2
    theme: Montpellier
    toc: yes
  ioslides_presentation:
    colortheme: dolphin
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    fonttheme: structurebold
    highlight: tango
    incremental: no
    keep_md: no
    slide_level: 2
    smaller: yes
    theme: cerulean
    toc: yes
    widescreen: yes
  pdf_document:
    fig_caption: yes
    highlight: zenburn
    toc: yes
    toc_depth: 3
  slidy_presentation:
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    highlight: tango
    incremental: no
    keep_md: no
    self_contained: yes
    slide_level: 2
    smaller: yes
    theme: cerulean
    toc: yes
    widescreen: yes
font-import: http://fonts.googleapis.com/css?family=Risque
font-family: Garamond
transition: linear
runtime: shiny_prerendered
---

<!--  
Here the parameters about the documents.
https://www.rstudio.com/wp-content/uploads/2015/03/rmarkdown-reference.pdf 
--> 


```{css, echo=FALSE}

```

<script language="JavaScript" type="text/javascript">
          
          function sizeTbl2(h,i) {
          var tbl = document.getElementById("section-" + i);
          tbl.style.display = h;
          }

</script>



<style>
.exo {
  border-radius: 5px;
  margin-top: 5px;
  margin-bottom: 5px;
  padding-top: 5px;
  padding-bottom: 5px;
  padding-left: 5px;
  padding-right: 5px;
  background-color: #FFEEFF;
  color: rgb(51, 51,153);
}
.tips {
       padding-top: 5px;
       padding-bottom: 5px;
       padding-left: 5px;
       padding-right: 5px;
       border: 1px dashed #2f6fab;
       background-color: #EEFFEE;
}
.solution {
            margin-top: 5px;
            margin-bottom: 5px;
            padding-top: 5px;
            padding-bottom: 5px;
            padding-left: 5px;
            padding-right: 5px;
            border: 1px dashed #FFFFFF;
            background-color: #EEEEFF;
            color: #0000BB;
            font-size: 11px;
}
</style>


```{r setup, include=FALSE}
library(learnr)
library(gradethis)
knitr::opts_chunk$set(echo=TRUE, eval=TRUE, cache=FALSE, message=FALSE, warning=FALSE, comment="")
gradethis::gradethis_setup()
```

<!--  
<img src="img/graph_with_r.jpeg"> 
![](img/graph_with_r.jpeg)
--> 

## Indexation

### Indexation ?

Indexer signifie récupérer certains éléments d'un objet, ici d'un vecteur. L'opérateur d'indexation est le crochet '[]', comme dans de nombreux autres langages. On l'utilisera sous la forme suivante:

```
# les_positions_ciblées sont les positions que l'on souhaite
# extraire du vecteur.
mon_vecteur[les_positions_ciblees]
```

Les positions ciblées (les_positions_ciblees) prendront la forme:
  
- Un vecteur d’indices (*i.e* de positions). Les positions sont comptées de $1$ à $n$ (avec $n$ la taille du vecteur).
- Ou un vecteur logique (vrai/faux).
- Ou les noms des positions ciblées dans le vecteur ci-celles ci sont nommées.

Nous allons présenter ces différents cas dans la suite.

## Indexation des vecteurs par position

L’argument de l'opérateur d'indexation sera ici un ensemble de positions au format numérique. 

**NB:** L'indexation dans R est **"one-based"**: **l'accès à la première position se fait avec la valeur d'index 1** (au contraire du langage Python pour lequel l'indexation est **"zero-based"**, on accède à la première position avec la valeur 0).

```{r vector12, exercise=TRUE}
set.seed(123) 
x <- sort(round(rnorm(10), digit=2), decreasing = T)
x
```
```{r vector12_1, exercise=TRUE, exercise.setup="vector12" }
x[1]
```

```{r vector12_2, exercise=TRUE, exercise.setup="vector12" }
x[1:5]
```

  On peut demander aussi toutes les positions hormis certaines (on utilise **-**).

```{r vector13, exercise=TRUE, exercise.setup="vector12"}
print(x[-2])
```

```{r vector13_1, exercise=TRUE, exercise.setup="vector12"}
print(x[-c(1,3)])
```

```{r vector13_2, exercise=TRUE, exercise.setup="vector12"}
print(x[-c(3:5)])
```


<div class="exo">
- Etant donnée le vecteur ci-dessous, utilisez l'opérateur d'indexation et :
  - Stockez dans la variable *a* la position 6 de x.
  - Stockez dans la variable *b* les position 5, 18 et 27.
  - Après avoir positionné le générateur de valeur aléatoire sur 123, tirez au hasard 10 position de x sans remise et placez les valeurs correspondantes dans la variable c.
  
```{r vector13a, exercise=TRUE}
set.seed(123)
x <- sort(sample(1:1000, size=100))
```

```{r vector13a-solution}
set.seed(123)
x <- sort(sample(1:1000, size=100))
a <- x[6]
b <- x[c(5, 18, 27)]
set.seed(123)
rnd_pos <- sample(1:length(x), size=10, replace=FALSE) 
c <- x[rnd_pos]
```

```{r vector13a-check}
gradethis::grade_result(
  pass_if(~ a == 39),
  pass_if(~ sum(b) == 411),
  pass_if(~ sum(c) == 4854)
)
```
</div>



## Indexation logique des vecteurs

Le principe est de passer à l'opérateur d'indexation un vecteur logique de même taille que le vecteur. **Seul les position vraies (TRUE), seront renvoyées/sélectionnées**.

```{r vector14, exercise=TRUE}
set.seed(123) 
x <- sort(round(rnorm(8), digit=2), decreasing = FALSE)
x
```

```{r vector14_1, exercise=TRUE, exercise.setup="vector14" }
x > 0
```

```{r vector14_2, exercise=TRUE, exercise.setup="vector14"}
x[x > 0]
```
Nous avons vu que l'on peut effectuer des opérations logiques avec les opérateurs "&" et "|". On peut donc aussi utiliser ces opérateurs pour tester les positions d'un vecteur sur la base d'un autre vecteur (si ils ont tous les deux mêmes tailles). 

Imaginons par exemple les coordonnées x et y de points en 2D. 


```{r vector15, exercise=TRUE}
set.seed(123)
# On crée des valeurs normalement distribuées 
# sur l'axe des x.
x <- rnorm(100)
# On ajoute un peu de bruit à x
# pour créer y.
y <- x + rnorm(100, mean=0.3, sd=0.4)
```

On peut visualiser ces points à l'aide de la fonction **plot()**.

```{r vector15a, exercise=TRUE, exercise.setup="vector15" }
# plot() permet de créer un nuage de points
plot(x, y)
# On ajoute une grille vertical/horizontal
grid()
# On ajoute une ligne vertical
# d'équation x=0
abline(v = 0, col="black")
# On ajoute une ligne horizontale
# d'équation y=0
abline(h= 0, col="black")
# On ajoute une ligne diagonale
# d'équation y=x
abline(a= 0, b=1, col="black")
```

On pourra par exemple facilement colorer les points sur l'axe des x et sur l'axe des y. Par exemple, ici, tous les points positifs en x et y.

```{r vector15b, exercise=TRUE, exercise.setup="vector15" }
plot(x, y)
grid()
abline(v = 0, col="black")
abline(h= 0, col="black")
abline(a= 0, b=1, col="black")

# Quelles sont les positions positives en x et y ?
pos <- x > 0 & y > 0
# points() permet de superposer des points
# à un diagramme existant.
# pch= le type de point
# col= la couleur
points(x[pos], y[pos], pch=16, col="red")
```


<div class="exo">
Terminez le code pour afficher en rouge les points pour lesquels les valeurs en x sont supérieures aux valeurs en y et en bleu celles pour lesquels y est supérieur à x.

```{r vector16, exercise=TRUE, eval=FALSE}
set.seed(123)
x <- rnorm(100)
y <- x + rnorm(100, mean=0.1, sd=0.3)
plot(x,y)
grid()
```

```{r vector16-solution}
set.seed(123)
x <- rnorm(100)
y <- x + rnorm(100, mean=0.1, sd=0.3)
plot(x,y)
grid()
abline(v = 0, col="black")
abline(h= 0, col="black")
abline(a= 0, b=1, col="black")
points(x[x > y], y[x > y], pch=16, col="red")
points(x[x < y], y[x < y], pch=16, col="blue")
```
</div>

<div class="exo">

Terminez le code pour afficher en rouge tous les points compris dans un cercle de centre 0 et de rayon 2 (pensez au théorème de Pythagore !). 

<div class="hint">
Pour élever à la puissance il faut utiliser l'opérateur `^`. On écrira par exemple `x ^ 2`.
</div>

```{r vector17, exercise=TRUE, eval=FALSE, fig.dim=c(5,5.5)}
set.seed(123)
x <- rnorm(100)
y <- rnorm(100)
plot(x,y)
```

```{r vector17-solution}
set.seed(123)
x <- rnorm(100)
y <- rnorm(100)
plot(x,y)
grid()
abline(v = 0, col="black")
abline(h= 0, col="black")
abline(a= 0, b=1, col="black")
pos <- x ^ 2 + y ^ 2 < 2
points(x[pos], y[pos], col="red", pch=16)
```
</div>

##  Indexation des vecteurs par les noms des positions

Les position d'un vecteur peuvent être nommées. Si elles le sont, on pourra extraire des position en passant leurs noms à l'opérateur d'indexation ('[]') sous la forme d'un objet vecteur.

Par exemple, un vecteur pourrait contenir l'expression de gènes analysés par QPCR. Chaque position serait alors le symbole du gène.

```{r vector18, exercise=TRUE}
gene_expression <- c(2, 8, 5.6, 10, 2.7)
names(gene_expression) <- c("CD3E", "CD4", "ZAP70", "PCNA", "BUB3")
gene_expression
```

On pourra donc naturellement rechercher l'expression des gènes dans le vecteur en utilisant les noms des gènes passés sous la forme d'un vecteur de chaînes de caractères.

```{r vector19, exercise=TRUE, exercise.setup="vector18"}
gene_expression[c("ZAP70", "CD3E", "CD4")]
```

##  Utilisation d'expressions régulières pour l'indexation

Reprenons ce même vecteur:

```{r vector19_1, exercise=TRUE, exercise.setup="vector18"}
gene_expression
```

On pourra éventuellement rechercher avec la fonction **grep()** les genes qui vérifient une expression régulière donnée. Les expressions régulières permettent d'effectuer **des recherches de motifs dans des chaînes de caractères**. La syntaxe founit un **ensemble d'opérateurs permettant de décrire le motif recherché**.

vous pouvez ci-dessous (bouton *View tips*) afficher un tableau avec les expressions régulières de base.

<a href="javascript:sizeTbl2('block', 'regexp')">View tips</a>|
<a href="javascript:sizeTbl2('none', 'regexp')">Hide tips</a>
<div id="regexp" class='solution' style="overflow:hidden; display:none">

|        Opérateur          |        Signification                                                      |
|------------------|---------------------------------------------------------------------------------------------|
| .                | Un caractère quelconque (.sauf 	\textbackslash n)                                             |
| $x*$               | 0 ou n fois le caractères $x$.                                                              |
| $x+$               | 1 ou n fois le caractère $x$.                                                               |
| $x\{n,m\}$         | Le caractère $x$ répété entre $n$ et $m$ fois.                                                  |
| $[a-z]$          | Une lettre minuscule (intervalle, ex: $[u-w]$).                                               |
| $[A-Z]$          | Une lettre majuscule (intervalle, ex: $[A-E]$).                                               |
| $[ABc]$          | $A$ ou $B$ ou $c$.                                                                                |
| [^ABab]   | Toute lettre différente de $a$ et $b$.                                                          |
| $(toto|yaya)$      | Une chaîne de caractères ou une autre.                                                      |
| \^{}             | Début de ligne.                                                                             |
| \$               | Fin de ligne.                                                                               |
| \\n | Retour à la ligne.                                                                          |
| \\t | Une tabulation.                                                                             |
| \\d | Un chiffre (**d**igit, $[0-9]$) .                                                      |
| \\w | Un caractère rencontré dans un mot (un alphanumérique ou le caractère '\_', [0-9a-zA-Z\_]). |
| \\s | Un blanc ([ \\t \\r \\n \\f]).             |
| \\D | Tout sauf un chiffre.                                                                       |
| \\W | Tout sauf un caractère rencontré dans un mot.                                               |
| \\S | Tout sauf un "espace blanc".                                                                         |
| \\   | Caractère d'échappement.                                                                    |
</div>


La fonction **grep()** renvoie, dans un vecteur de chaînes de caractères **les positions vérifiant une expression régulière**. Par exemple, ici, on peut chercher avec **grep()** toutes les positions contenant un gène dont le nom commence par "CD".

**NB**: Par defaut **grep()** ne comprend pas certains opérateurs. L'utilisation, ci-dessous de l'argument *perl = TRUE* permet d'étendre ses capacités à tous les opérateurs. Dans le doute, il est donc préferable de positionner cet argument sur TRUE.

```{r vector20, exercise=TRUE, eval=F, exercise.setup="vector18"}
pos <- grep(pattern = "^CD", x = names(gene_expression), perl = TRUE)
pos
```

On peut ensuite simplement extraire ces positions.

```{r vector20_1, exercise=TRUE, eval=F, exercise.setup="vector20"}
gene_expression[pos]
```

<div class="exo">

Dans l'exemple ci-dessous, extraire les valeurs d'expression pour les gènes commençant par 'CD' et finissant par une valeur numérique. Si vous débutez en expression régulière, 

```{r vector21, exercise=TRUE, eval=FALSE}
gene_expression <- c(2, 8, 5.6, 10, 2.7, 4, 9, 12)
names(gene_expression) <- c("CD3E", "CD4", "ZAP70", "PCNA", "BUB3", "CDC42BPA", "CDK11B", "CD52")
```

```{r vector21-solution}
gene_expression <- c(2, 8, 5.6, 10, 2.7, 4, 9, 12)
names(gene_expression) <- c("CD3E", "CD4", "ZAP70", "PCNA", "BUB3", "CDC42BPA", "CDK11B", "CD52")
pos <- grep(pattern = "^CD.*[0-9]$", x = names(gene_expression), perl = TRUE)
gene_expression[pos]
```
</div>

<div class="exo">
Dans l'exemple ci-dessous, extraire les valeurs d'expression pour les gènes commençant par 'CD' et finissant par un caractère alphanumérique.

```{r vector21b, exercise=TRUE, eval=FALSE}
gene_expression <- c(2, 8, 5.6, 10, 2.7, 4, 9, 12)
names(gene_expression) <- c("CD3E", "CD4", "ZAP70", "PCNA", "BUB3", "CDC42BPA", "CDK11B", "CD52")
```

```{r vector21b-solution}
gene_expression <- c(2, 8, 5.6, 10, 2.7, 4, 9, 12)
names(gene_expression) <- c("CD3E", "CD4", "ZAP70", "PCNA", "BUB3", "CDC42BPA", "CDK11B", "CD52")
pos <- grep(pattern = "^CD.*[0-9A-Za-z]$", x = names(gene_expression), perl = TRUE)
gene_expression[pos]
```
</div>

<h1> Fin de la section  </h1>
