---
# theme
# “bootstrap”, “cerulean”, “cosmo”, “darkly”, “flatly”, “journal”, “lumen”, “paper”, 
# “readable”, “sandstone”, “simplex”, “spacelab”, “united”, “yeti”
# highlight: `default`, `tango`, `pygments`, `kate`, `monochrome`, `espresso`, `zenburn`, 
# `haddock`, `breezedark`, `textmate`, `arrow`, or `rstudio` or a file with extension `.theme`.
# Estimated time = 2h
title: "Les objets data.frame"
output:
  learnr::tutorial:
    theme: default
    highlight: default
    fig_caption: yes
    self_contained: true
    toc: yes
    toc_float: 
      toc_collapsed: false
    toc_depth: 4
    number_sections: false
    progressive: true
  html_document:
    theme: cosmo
    fig_caption: yes
    self_contained: yes
    toc: yes
    toc_float: 
      toc_collapsed: false
    toc_depth: 3
    number_section: true
  beamer_presentation:
    colortheme: dolphin
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    fonttheme: structurebold
    highlight: tango
    incremental: no
    keep_tex: no
    slide_level: 2
    theme: Montpellier
    toc: yes
  ioslides_presentation:
    colortheme: dolphin
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    fonttheme: structurebold
    highlight: tango
    incremental: no
    keep_md: no
    slide_level: 2
    smaller: yes
    theme: cerulean
    toc: yes
    widescreen: yes
  pdf_document:
    fig_caption: yes
    highlight: zenburn
    toc: yes
    toc_depth: 3
  slidy_presentation:
    fig_caption: yes
    fig_height: 4
    fig_width: 6
    highlight: tango
    incremental: no
    keep_md: no
    self_contained: yes
    slide_level: 2
    smaller: yes
    theme: cerulean
    toc: yes
    widescreen: yes
font-import: http://fonts.googleapis.com/css?family=Risque
font-family: Garamond
transition: linear
runtime: shiny_prerendered
---

<!--  
Here the parameters about the documents.
https://www.rstudio.com/wp-content/uploads/2015/03/rmarkdown-reference.pdf 
--> 


```{css, echo=FALSE}

```

<script language="JavaScript" type="text/javascript">
          
          function sizeTbl2(h,i) {
          var tbl = document.getElementById("section-" + i);
          tbl.style.display = h;
          }

</script>



<style>
.exo {
  border-radius: 5px;
  margin-top: 5px;
  margin-bottom: 5px;
  padding-top: 5px;
  padding-bottom: 5px;
  padding-left: 5px;
  padding-right: 5px;
  background-color: #fcede3;
  color: rgb(51, 51,153);
}
.tips {
       padding-top: 5px;
       padding-bottom: 5px;
       padding-left: 5px;
       padding-right: 5px;
       border: 1px dashed #2f6fab;
       background-color: #EEFFEE;
}
.solution {
            margin-top: 5px;
            margin-bottom: 5px;
            padding-top: 5px;
            padding-bottom: 5px;
            padding-left: 5px;
            padding-right: 5px;
            border: 1px dashed #FFFFFF;
            background-color: #EEEEFF;
            color: #0000BB;
            font-size: 11px;
}
</style>


```{r setup, include=FALSE}
library(learnr)
library(gradethis)
knitr::opts_chunk$set(echo=TRUE, eval=TRUE, cache=FALSE, message=FALSE, warning=FALSE, comment="")
gradethis::gradethis_setup()
```


<!--  
<img src="img/graph_with_r.jpeg"> 
![](img/graph_with_r.jpeg)
--> 


## Creation

### The data.frame() Function

The `data.frame` object is truly central in R. This object allows to store various information about objects under study (patients, biological samples, experiments...). In short, a data.frame can be viewed as a matrix where **each column can have a different and unique mode (*character, numeric, boolean*, etc.)**.

You can create a dataframe using the `data.frame()` function. In the example below, the `note` column is numeric, the `diploma` column is a factor, the `age` column is numeric, and the `name` column is a character string.



```{r dataframe0, exercise=TRUE}
set.seed(123)
df_1 <- data.frame(
  note = round(runif(n = 6, min = 0, max = 20), 2),
  diploma = as.factor(sample(c("B.Sc", "M.Sc"), size = 6, replace = TRUE )),
  age = sample(18:25, size=6, replace=TRUE),
  name = c("Miescher", "Watson", "Crick", "Chargaff", "Levene", "Kossel")
)
df_1
```


### Exercice

<div class="exo">

- Create a `data.frame` named `my_df` containing the vectors and factors below (the columns will be named "col_A", "col_B", "col_C", and "col_D", respectively).

```{r df1, exercise=TRUE}
set.seed(123)
c_a <- round(rnorm(10), 2)
c_b <- 1:10
c_c <- letters[1:10]
c_d <- factor(sort(rep(c("M", "F"), 5)), ordered=TRUE)
```

```{r df1-solution}
set.seed(123)
c_a <- round(rnorm(10), 2)
c_b <- 1:10
c_c <- letters[1:10]
c_d <- factor(sort(rep(c("M", "F"), 5)), ordered=TRUE)
my_df <- data.frame(col_A=c_a, col_B=c_b, col_C=c_c, col_D=c_d)
```

```{r df1-check}
gradethis::grade_result(
  pass_if(~ round(sum(my_df$col_A), 0) == 2),
  pass_if(~ round(sum(my_df$col_B), 55) == 2),
  pass_if(~ all(my_df$col_C == letters[1:10])),
  pass_if(~ all(levels(my_df$col_D)) == c("F", "M"))
)
```
</div>

### The as.data.frame() and as.matrix() Functions

You can **convert a `matrix` object into a `data.frame` object using the `as.data.frame()` function**. This allows you to present an object in the correct format for a function that requires a `data.frame` instead of a matrix.

```{r dataframe02,  exercise=TRUE}
mat <- matrix(data=round(rnorm(20), 2), 
            ncol = 5)
df_2 <- as.data.frame(mat)
print(df_2)
```

Similarly, you can **convert a `data.frame` into a matrix using `as.matrix()`.**

<div class="exo">

- Since `data.frame` objects may contain elements other than numeric values, the `hist()` function will not accept them by default. Convert the following `data.frame` into a matrix (object $mat$) and pass this matrix to the `hist()` function to construct the corresponding histogram.

```{r dataframe03, exercise=TRUE}
set.seed(123)
df <- data.frame(A=rnorm(10000, mean=0, sd=1),
                 B=rnorm(10000, mean=5, sd=2),
                 C=rnorm(10000, mean=10, sd=1)
                 )

# L'instruction suivante renvoie une erreur
# hist(df)
# Error in hist.default(df) : 'x' must be numeric

# Il faut donc convertir
```

```{r dataframe03-solution}
set.seed(123)
df <- data.frame(A=rnorm(10000, mean=0, sd=1),
                 B=rnorm(10000, mean=5, sd=2),
                 C=rnorm(10000, mean=10, sd=1)
                 )

mat <- as.matrix(df)
hist(mat, breaks=100)
```

```{r dataframe03-check }
gradethis::grade_result_strict(
  pass_if(~ round(sum(mat), 0) == 149723), 
  pass_if(~ is.matrix(mat) == TRUE)
)
```
</div>

## Datasets

You can also use one of the many datasets provided by the R software, many of which are `data.frame` objects.

<div class="exo">

- List the available demonstration datasets using the `data()` function without any arguments.

```{r data, exercise=TRUE }

```

```{r data-solution}
data()
```

</div>

Below, we load the `ChickWeight` dataset (*chicken weights*), which contains weight measurements used to evaluate the effectiveness of various dietary supplements on the growth rate of chickens. This dataset is provided by the `datasets` library. Therefore, it is necessary to load this library into memory (`library(datasets)`) to access the dataset. The dataset is explicitly loaded with `data(ChickWeight)`.

```{r dataframe1, exercise=TRUE}
# On charge/active la librairie datasets
library(datasets)
# On charge le jeu de donnée chickwts
data(ChickWeight)
# Voilà ses premières lignes
head(ChickWeight)
```



## Row and Column Names

As with matrices, **the column names of a `data.frame` can be obtained with `colnames()`** (and row names with `rownames()`).

```{r dataframe2, exercise=TRUE, setup.exercise="dataframe1"}
colnames(ChickWeight)
head(rownames(ChickWeight))
```

As with matrices, you can change the names of the columns or rows.

<div class="exo">

- Use the `paste0()` function to change the row names of `ChickWeight` and rename them "line_1", "line_2", "line_3", ...


```{r linedf, exercise=TRUE, exercise.setup="dataframe1" }
ChickWeight
```

```{r linedf-solution }
rownames(ChickWeight) <- paste0("line_", 1:nrow(ChickWeight))
```

```{r linedf-check }
gradethis::grade_result(
  pass_if(~ all(head(rownames(ChickWeight)) == c("line_1", "line_2", "line_3", "line_4", "line_5", "line_6")))
)
```

</div>

## Indexing

A `data.frame` can be indexed by querying rows and columns, just like a matrix. As with a matrix, you can use **position indexing, logical indexing, and name-based indexing**. However, the `data.frame` also has an additional indexing operator, the dollar sign `$`.

Type `ChickWeight$` and press the [<TAB> key](https://www.computerhope.com/jargon/t/tab-key.htm) to display the contents of the `feed` column. Do the same for the `weight` column.


```{r dataframe3, exercise=TRUE, setup.exercise="dataframe1"}
# Place your cursor after 'chickwts', type '$' and 
# press the tab key (<TAB>)
ChickWeight
```

Using logical indexing as we have seen with vectors and matrices, we can write the following to select rows in the `data.frame`:


```{r idx1, exercise=TRUE, setup.exercise="dataframe1"}
ChickWeight[ChickWeight$weight > 300, ]
```

Or even:

```{r idx2, exercise=TRUE, setup.exercise="dataframe1"}
ChickWeight[ChickWeight$weight > 300 & ChickWeight$Diet == 2,  ]
```

<div class="exo">

Using the **$** operator and indexing, extract the rows of `ChickWeight` where:

- The type of food (`Diet`) is 1 or 2. Store the result in `a`.
- The weight of the chicken is between 300 and 350 (inclusive). Store the result in `b`.
- The type of food (`Diet`) is 1 or 2 for an age (`Time` column) of 16 days. Store the result in `c`.

```{r chickindex, exercise=TRUE, setup.exercise="dataframe1"}
ChickWeight
```

```{r chickindex-solution}
a <- ChickWeight[ChickWeight$Diet == 1 | ChickWeight$Diet == 2, ]
b <- ChickWeight[ChickWeight$weight >= 300 & ChickWeight$weight <= 350, ]
c <- ChickWeight[(ChickWeight$Diet == 1 | ChickWeight$Diet == 2) & ChickWeight$Time == 16, ]
```

```{r chickindex-check }
gradethis::grade_result_strict(
  pass_if(~ nrow(a) == 340),
  pass_if(~ nrow(b) == 12),
  pass_if(~ nrow(c) == 27)
)
```
</div>

## Adding Columns

You can easily add a column to a **data.frame**, especially by using the *$* operator.

```{r addcol, exercise=TRUE, setup.exercise="dataframe1"}
ChickWeight$genotype <- sample(c("wild_type", "transgenic"), size=nrow(ChickWeight), replace=TRUE)
head(ChickWeight)                        
```


<div class="exo">

Add to the `ChickWeight` table a column named:

- 'genotype' containing NA values.
- 'source' containing "anonymous".
- 'location' containing random values chosen from 'box_1' or 'box_2'.

```{r chickaddcol, exercise=TRUE, setup.exercise="dataframe1"}
ChickWeight
```

```{r chickaddcol-solution}
ChickWeight$genotype <- NA
ChickWeight$source <- "anonymous"
ChickWeight$location <- sample(c('box_1', 'box_2'), nrow(ChickWeight), replace = TRUE)
```

```{r chickaddcol-check }
gradethis::grade_result_strict(
  pass_if(~ all(is.na(ChickWeight$genotype)) == TRUE),
  pass_if(~ all(ChickWeight$source == "anonymous") == TRUE),
  pass_if(~ all(ChickWeight$location %in% c('box_1', 'box_2')) == TRUE)
)
```
</div>

## 'wide' vs 'long'

**For many applications, especially for graphics with the `ggplot2` library** (which is indispensable, and we will discuss it further later), it is necessary to change the representation of data in a `data.frame` or a matrix.

Let’s take the following matrix. It is initially in the so-called **"wide" format**. At this stage, there is no formal way to identify this (you will see the difference with the **long** format shortly).


```{r longlarge0, exercise=TRUE, eval=TRUE}
# We populate a matrix with data
# generated from a normal distribution (n=20). 
# The matrix will contain 5 columns.
set.seed(123)
m <- matrix(data=round(rnorm(20), 2), 
            ncol = 5)
colnames(m) <- LETTERS[1:5]
rownames(m) <- letters[1:4]
print(m)
```

The `melt()` function from the `reshape2` library allows transforming (or "melting") it into the **long** format. The number of rows is now the product of the number of rows and columns of the input object.


```{r largelong1, exercise=TRUE, eval=TRUE, exercise.setup="longlarge0"}
library(reshape2)
m_melted <- melt(m)
dim(m_melted)
head(m_melted)
```

Note that the **result is different if the `matrix` is first converted into a `data.frame`** using `as.data.frame()`. 


```{r largelong2, exercise=TRUE, eval=TRUE, exercise.setup="longlarge0"}
library(reshape2)
df_melted <- melt(as.data.frame(m))
dim(df_melted)
head(df_melted)
```

## Applying Functions

As with matrices, you can apply functions to the rows and columns of a `data.frame` using the `apply()` function (*see* the chapter on matrices). 

```{r lapply2, exercise=TRUE, exercise.setup="largelong2" }
apply(df_melted, 2, length)
```

### End of the section

Thank you for following this tutorial.
